```python
# Initialize Otter
import otter
grader = otter.Notebook("lab01.ipynb")
```

# CE 93: Lab Assignment 01

You must submit the lab to Gradescope by the due date. You will submit the zip file produced by running the final cell of the assignment.

## About this Lab
The objective of this assignment is to revisit fundamental Python commands. As outlined in the syllabus, Data 8 is a prerequisite for CE 93. More specifically, basic programming experience is required to succeed in this course. While CE 93 will extend your Python proficiency, we won't cover the basics, assuming that students already possess this foundational knowledge. If you find it challenging to provide correct answers to more than 70% of the lab questions, it is advisable to wait until you have completed Data 8 (or an equivalent course) before enrolling in CE 93.

## Instructions 
**Run the first cell, Initialize Otter**, to import the autograder and submission exporter.

Throughout the assignment, replace `...` with your answers. We use `...` as a placeholder and theses should be deleted and replaced with your answers.

Any part listed as a "<font color='red'>**Question**</font>" should be answered to receive credit.

**Please save your work after every question!**

To read the documentation on a Python function, you can type `help()` and add the function name between parentheses.

**Run the cell below**, to import the required modules.

```python
# Please run this cell, and do not modify the contents
import math
import numpy as np
import scipy
import pandas as pd
import statistics as stats
import cmath
import re
import matplotlib.pyplot as plt
import matplotlib.image as mpimg
import hashlib
import ipywidgets as widgets
from ipywidgets import FileUpload
from IPython.display import display
from PIL import Image
import os
import resources
from resources.utils import *

def get_hash(num):
    """Helper function for assessing correctness"""
    return hashlib.md5(str(num).encode()).hexdigest()
```

### Python Basics

<font color='red'>**Question 1.0.**</font> What does a hashtag (`#`) represent in Python syntax? Assign your answer to the variable `q1_0` as a string. (0.5 pts)

A. Assignment \
B. Function \
C. Comment \
D. Selection \
E. Iteration

Answer in the next cell. Your answer should be a string, e.g., `"A"`, `"B"`, etc. \
For example, if you selected `A. Assignment`, then simply replace `...` in the cell below by `"A"`.
Remember to put quotes around your answer choice.

```python
# ANSWER CELL
q1_0 = ...
q1_0
```

```python
grader.check("q1.0")
```

### Mathematical Operations

An arithmetic **operation** is either addition, subtraction, multiplication, division, or exponentiation (i.e., power). An arithmetic **operator** is a symbol that Python has reserved to perform one of the aforementioned operations.

| Operation       | Mathematical Notation | Python Operator | Python Example | Output    |
|:----------------|:---------------------:|:---------------:|:---------------|:----------|
| Addition        | $a+b$                 | `+`             | `3 + 1`        | `5`       |
| Subtraction     | $a-b$                 | `-`             | `3 - 2`        | `1`       |
| Multiplication  | $a\times b$           | `*`             | `3 * 2`        | `6`       |
| Division        | $\dfrac{a}{b}$        | `/`             | `3 / 2`        | `1.5`     |
| Exponentiation  | $a^b$                 | `**`            | `3 ** 2`       | `9`       |

Python obeys the same order of operations that you learned in school: powers are executed before multiplication and division, which are executed before addition and subtraction. Parentheses, `()`, can also be used in Python to group together smaller expressions and supersede the standard order of operations. The order in which Python executes these operations is shown in the table below.

| Precedence      | Operation                      | Rule                                       | 
|:----------------|:------------------------------ | :----------------------------------------- |
| First (Highest) | Parentheses: `( )`             | Evaluated starting with the innermost pair |
| Second          | Exponentiation: `**`           | Evaluated from left to right               | 
| Third           | Unary positive, negative: `+x`, `-x` | Evaluated from left to right         |
| Fourth          | Multiplication, division: `*`, `/`   | Evaluated from left to right         |
| Fifth           | Addition, subtraction: `+`, `-`      | Evaluated from left to right         |

<font color='red'>**Question 2.0.**</font> What is the result of this Python expression? Feel free to run it to check your answer, but try to think it through first. Assign your answer to the variable `q2_0` as a string. (0.5 pts)

```python
2 * 3 ** 2
```

A. 9 \
B. 12 \
C. 18 \
D. 36 \
E. `Error`

Answer in the next cell. Your answer should be a string, e.g., `"A"`, `"B"`, etc. \
Assign your answer to the given variable.
Remember to put quotes around your answer choice.

```python
# ANSWER CELL
q2_0 = ...
q2_0
```

```python
grader.check("q2.0")
```

<font color='red'>**Question 2.1.**</font> Construct a one-line Python expression that computes the value of the following mathematical expression. Assign the expression to the variable `q2_1`. Your answer should be a mathematical expression, not the value of the expression. (0.5 pts)

$$3+\frac{1}{3}$$

```python
# ANSWER CELL
q2_1 = ...
q2_1
```

```python
grader.check("q2.1")
```

Python has many basic trigonometric, logarithmic, and other arithmetic functions like `sin()`, `cos()`, `tan()`, `asin()`, `acos()`, `atan()`, `exp()`, `log()`, `log10()` and `sqrt()` stored in a module called **`math`**. You can read about all the functions provided by the **`math`** module [here](https://docs.python.org/3/library/math.html).

The way we use these mathematical functions (and in general, any function from a module) is `module.function(x)`, where:
* `module` is the module's name or its given shortened name; e.g., `math`
* `.function` is the function available within that module that you want to use; e.g., `.log10`
* `(x)` includes inside the parentheses the input(s) that your are trying to evaluate using the function; e.g., `(100)`.

**Example:** 

```python
>>> math.log10(100)
2.0
```

<font color='red'>**Question 2.2.**</font>  Define variables $x=2$ and $y=5$. Then, construct a one-line Python expression that computes the value of the following mathematical expression. Assign the expression to the variable `q2_2`. (0.5 pts)

$$e^{x^y}$$

Your Python expression should give appropriate values regardless of the values of $x$ and $y$ (i.e. you should use the variable names instead of their numerical values).

*Hint:* Take a look at `math`'s exponential function documentation [here](https://docs.python.org/3/library/math.html#power-and-logarithmic-functions).

```python
# ANSWER CELL
x = ...
y = ...
q2_2 = ...
q2_2
```

```python
grader.check("q2.2")
```

### Variables and Assignment

**Variables** are used to store values in Python using an *assignment* statement. In an *assignment* statement, a name is followed by `=`, which is followed by any expression. That's why <code>=</code> is referred to as the assignment operator.

**Example:**

```python
a = 5
b = a
```

The value of the expression to the right of `=` is *assigned* to the variable name to the left of `=`. So, in the example above, the value 5 is assigned to the variable `a`. After executing this line, the variable `a` behaves like the value 5 in subsequent calculations, until the variable is changed or deleted.

When a variable is assigned, it has no memory of how it was assigned. We assigned variable <code>b = a</code>. However, changing the value of <code>a</code> will not automatically change the value of <code>b</code>. To change the value of <code>b</code>, you have to **reassign it**.

<font color='red'>**Question 3.0.**</font> What will be displayed after running this Python code? Feel free to run it to check your answer, but try to think it through first. Assign your answer to the variable `q3_0` as a string. (0.5 pts)

```python
a = [1, 2, 3]
b = a
a = [4, 5, 6]
print(b)
```

A. `[1, 2, 3]` \
B. `b` \
C. `[4, 5, 6]` \
D. `a` \
E. No output

Answer in the next cell. Your answer should be a string, e.g., `"A"`, `"B"`, etc. \
Assign your answer to the given variable.
Remember to put quotes around your answer choice.

```python
# ANSWER CELL
q3_0 = ...
q3_0
```

```python
grader.check("q3.0")
```

<font color='red'>**Question 3.1.**</font> What will be displayed after running this Python code? Feel free to run it to check your answer, but try to think it through first. Assign your answer to the variable `q3_1` as a string. (0.5 pts)

```python
a = 10
b = a
a = 15
```

A. `10` \
B. `15` \
C. `b` \
D. `a` \
E. No output

Answer in the next cell. Your answer should be a string, e.g., `"A"`, `"B"`, etc. \
Assign your answer to the given variable.
Remember to put quotes around your answer choice.

```python
# ANSWER CELL
q3_1 = ...
q3_1
```

```python
grader.check("q3.1")
```

The statement <code>a = 5</code> is not asserting that <code>a</code> is already equal to 5, as we might expect in mathematical notation. Rather, this line of code defines the variable <code>a</code> to hereafter have the value 5.

Although it is perfectly valid to say $5 = a$ in mathematics, `5 = a` will raise an error in Python.

When trying to run `5 = a`, you are telling Python to define a variable named `5` and assign it the value of the variable `a`. However, `5` is not a valid variable name. That's why Python raises a `SyntaxError`. There are some restrictions on the names variables can take in Python, which are listed hereafter:
* Variable names must start with a letter or underscore `_` but cannot start with a number
* Variables can only contain alphanumeric characters (letters and numbers) as well as underscores `_` 
* Numbers are allowed after the first character of the variable name
* No spaces or math symbols/punctuation within a variable name are permitted
* Variable names are case-sensitive (e.g., `a` and `A` will be considered different variables)

So, variable names must start with a letter, but can contain both letters and numbers after the first character. A name cannot contain a space; instead, it is common to use an underscore `_` character to replace each space. For example, `1dice_roll = 5` will raise a `SyntaxError` because a name cannot start with a number. Instead, you can write `dice_roll_1 = 5`, and then `dice_roll_1` will be assigned the value $5$.

<font color='red'>**Question 3.2.**</font>  What is the result of this Python expression? Feel free to run it to check your answer, but try to think it through first. Assign your answer to the variable `q3_2` as a string. (0.5 pts)

```python
2 = 3
```

A. `True` \
B. `False` \
C. `NaN` \
D. `NA` \
E. `SyntaxError`

Answer in the next cell. Your answer should be a string, e.g., `"A"`, `"B"`, etc. \
Assign your answer to the given variable.
Remember to put quotes around your answer choice.

```python
# ANSWER CELL
q3_2 = ...
q3_2
```

```python
grader.check("q3.2")
```

<font color='red'>**Question 3.3.**</font> Which of the following is (are) NOT valid variable name(s)? Add ALL that apply to the variable `q3_3`. (1 pt)

A. `Data` \
B. `Data1` \
C. `1Data` \
D. `^_^` \
E. `@potus`\
F. `Data 1`\
G. `Data_1`

Answer in the next cell. Add each selected choice as a string and separate each two answer choices by a comma. For example, if you want to select `"A"` and `"B"`, your answer should be `"A", "B"`.\
Assign your answer to the given variable.
Remember to put quotes around each answer choice.

```python
# ANSWER CELL
q3_3 = ...
q3_3
```

```python
grader.check("q3.3")
```

### Logical Expressions and Operators

**Logical expressions** are statements that can either be true or false. Logical expressions are used to pose questions to Python. For example, running `2 < 3` is equivalent to asking , "*Is 2 less than 3?*" Since this statement is true, Python will return the value `True`. Otherwise, if the statement is false, Python will return `False`. The returned value is a data type which is known as **boolean** (or `bool`), which has the built-in values `True` and `False`.

Logical expressions can be applied on values, variables, or both. Python includes a variety of comparison operators (operators to perform comparisons). 

| Comparison         | Python Operator | True example | False Example |
|:-------------------|:---------------:|:-------------|:--------------|
| Less than          | `<`             | `2 < 3`      | `2 < 2`       |
| Greater than       | `>`             | `3 > 2`      | `3 > 3`       |
| Less than or equal | `<=`            | `2 <= 2`     | `3 <= 2`      |
| Greater or equal   | `>=`            | `3 >= 3`     | `2 >= 3`      |
| Equal              | `==`            | `3 == 3`     | `3 == 2`      |
| Not equal          | `!=`            | `3 != 2`     | `2 != 2`      |

<font color='red'>**Question 4.0.**</font> Recall that you have previously defined `x = 2`. What is the result of this Python expression? Feel free to run it to check your answer, but try to think it through first. Assign your answer to the variable `q4_0` as a string. (0.5 pts)

```python
x != 3
```

A. `True` \
B. `False` \
C. `NaN` \
D. `NA` \
E. `SyntaxError`

Answer in the next cell. Your answer should be a string, e.g., `"A"`, `"B"`, etc. \
Assign your answer to the given variable.
Remember to put quotes around your answer choice.

```python
# ANSWER CELL
q4_0 = ...
q4_0
```

```python
grader.check("q4.0")
```

Two consecutive equal marks (`==`) are used to check whether two expressions have the same value. For example, if we define `dice_roll_1 = 5` and then run `dice_roll_1 == 5`, this will return `True` because we have assigned the value $5$ to `dice_roll_1`.

<font color='red'>**Question 4.1.**</font>  What is the result of this Python expression? Feel free to run it to check your answer, but try to think it through first. Assign your answer to the variable `q4_1` as a string. (0.5 pts)

```python
x == 3
```

A. `True` \
B. `False` \
C. `NaN` \
D. `NA` \
E. `SyntaxError`

Answer in the next cell. Your answer should be a string, e.g., `"A"`, `"B"`, etc. \
Assign your answer to the given variable.
Remember to put quotes around your answer choice.

```python
# ANSWER CELL
q4_1 = ...
q4_1
```

```python
grader.check("q4.1")
```

In Python and many other programming languages, `True` is equivalent to 1 and `False` is equivalent to 0. So, you can perform arithmetic operations on `bool` data types. Because `True` is equal to 1 and `False` is equal to 0, adding logical expressions together is a quick way to count the number of `True` expressions out of many logical expressions.

<font color='red'>**Question 4.2.**</font> What is the result of this Python expression? Feel free to run it to check your answer, but try to think it through first. (0.5 pts)

```python
(5 > 4) + 2 * (3 != 10) - (7 < 10)
```

Assign your answer to the variable `q4_2` in the next cell. Do NOT put quotes around your answer. Your answer should be a numeric value.

```python
# ANSWER CELL
q4_2 = ...
q4_2
```

```python
grader.check("q4.2")
```

In most of the previous examples, we did one comparison at a time. **Logical operators** allow us to combine together multiple comparisons. There are three common logical operators: AND, OR, NOT. Operator symbols vary with different programming languages. In Python they are `and`, `or`, and `not`.

| Operator | Description | Examples |
| :---:    | :---:       | :---:    | 
| `and`    | `True` if **all** expressions are `True` <br> `False` otherwise| `(5 > 4) and (5 > 3)`<br>`(5 > 4) and (5 < 3)`|
| `or`     | `True` if **at least one** expression is `True` <br> `False` otherwise| `(5 > 4) or (5 < 3)`<br>`(5 < 4) or (5 < 3)`|
| `not`    | Reverses the result, `True` if the expression is `False` <br> `False` if the expression is `True` |`not(5 < 4)`<br>`not(5 > 4)`| 

<font color='red'>**Question 4.3.**</font> What is the result of this Python expression? Feel free to run it to check your answer, but try to think it through first. Assign your answer to the variable `q4_3` as a string. (0.5 pts)

```python
(5 > 4 or 2 == 3)  and  (3 != 10 and 7 < 10)
```

A. `True` \
B. `False` \
C. `NaN` \
D. `NA` \
E. `SyntaxError`

Answer in the next cell. Your answer should be a string, e.g., `"A"`, `"B"`, etc. \
Assign your answer to the given variable.
Remember to put quotes around your answer choice.

```python
# ANSWER CELL
q4_3 = ...
q4_3
```

```python
grader.check("q4.3")
```

### Strings

A string is a sequence of characters, such as `"Hello World"`. Strings are surrounded by either single `' '` or double quotation marks `" "`. A string can include letters, digits, punctuation, spaces, and other valid symbols. They can represent anything from a single character to a word, a sentence, or even an entire book! Strings can also be assigned to variables.

Numbers can also be expressed as strings. For example, `z = "93"` means that `z` is assigned the string 93 not the number 93.

Strings can be concatenated (glued together) with the `+` operator. This allows you to combine multiple strings into a single string. So, if we define `z = '93'` and then run `"Welcome " + "to " + "CIVENG" + z`, this will return the string `"Welcome to CIVENG93"`. Notice how spaces were included after "Welcome" and after "to", but not after "CIVENG".

<font color='red'>**Question 5.0.**</font> What is the result of this Python expression? Feel free to run it to check your answer, but try to think it through first. Assign your answer to the variable `q5_0` as a string. (0.5 pts)

```python
'2' + '222'
```

A. `2222` \
B. `224` \
C. `222` \
D. `'2222'` \
E. `'224'`

Answer in the next cell. Your answer should be a string, e.g., `"A"`, `"B"`, etc. \
Assign your answer to the given variable.
Remember to put quotes around your answer choice.

```python
# ANSWER CELL
q5_0 = ...
q5_0
```

```python
grader.check("q5.0")
```

In Python, if you try to concatenate a string with a numeric value, you will get a `TypeError`.

```python
>>> course_code = 'CE'
>>> course_number = 93
>>> course_code + course_number
```
```
---------------------------------------------------------------------------
TypeError                                 Traceback (most recent call last)

TypeError: can only concatenate str (not "int") to str
```

The correct way to do it is to convert the numeric value to string first, and then concatenate them. This can be achieved using the built-in function `str()`. For example `str(1)` will convert the number 1, to the string '1'. This is particularly useful when you have calculated a numeric variable and want to print it out with a string. In this case, you can convert the numeric variable to a string and then concatenate it with the string you are trying to print out.

```python
>>> course_code = 'CE'
>>> course_number = 93
>>> course_code + str(course_number)
'CE93'
```

A more elegant way to combine numeric data with strings is known as formatted strings. Prefix the string with the letter `f` and then any variables, which could be strings, numeric, or other are included between `{ }`.

```python
>>> course_code = 'CE'
>>> course_number = 93
>>> f'Welcome to {course_code}{course_number}!'
'Welcome to CE93!'
```

<font color='red'>**Question 5.1.**</font> Consider the Python code below. Which of the options below could replace `...` to get the following output: `'The average rainfall is 0.925 inch'`? Add ALL that apply to the variable `q5_1`. (1 pt)

```python
rain_day1 = 1.1
rain_day2 = 0.75
...
```

A. `'The average rainfall is ' + str((rain_day1 + rain_day2) / 2) + ' inch'` \
B. `'The average rainfall is ' + (rain_day1 + rain_day2) / 2 + ' inch'` \
C. `'The average rainfall is {(rain_day1 + rain_day2) / 2} inch'` \
D. `f'The average rainfall is {(rain_day1 + rain_day2) / 2} inch'` \
E. `f'The average rainfall is (rain_day1 + rain_day2) / 2 inch'`

Answer in the next cell. Add each selected choice as a string and separate each two answer choices by a comma. For example, if you want to select `"A"` and `"B"`, your answer should be `"A", "B"`.
Assign your answer to the given variable. Remember to put quotes around each answer choice.

```python
# ANSWER CELL
q5_1 = ...
q5_1
```

```python
grader.check("q5.1")
```

### Lists and NumPy Arrays

A `list` is a more versatile data structure that is a collection of one or more items. Lists can hold various data types or a combination of different data types, such as strings and numeric data. Items in a list are enclosed within square brackets `[]` and separated by commas `,`.

**Examples:**

```python
>>> numeric_list = [1, 2, 3]
>>> numeric_list
[1, 2, 3]

>>> text_list = ['Welcome', 'to', 'CE93!']
>>> text_list
['Welcome', 'to', 'CE93!']

>>> mixed_list = ['Welcome', 'to', 'CE', 93, '!']
>>> mixed_list
['Welcome', 'to', 'CE', 93, '!']
```
Lists are fundamental in Python and serve as a powerful tool for storing and manipulating collections of data.

We can use indexing to access an element from a list. To access an element from a list, we use a pair of square brackets. Indexing in Python starts at 0, which means that the first element has an index of 0, the second element has an index of 1, and so on.

```python
>>> numeric_list = [12, 52, 32]
>>> numeric_list[0]
12
```

You can also use negative indexes. Negative indexes simply means counting from the end of the list. So, -1 is the index of the last element, -2 is the index of the second-to-last element, and so on.

```python
>>> numeric_list = [12, 52, 32]
>>> numeric_list[-1]
32
```

<font color='red'>**Question 6.0.**</font> What will be displayed after running this Python code? Feel free to run it to check your answer, but try to think it through first. (0.5 pts)

```python
a = [383, 165, 273, 112, 937, 11, 64, 112, 455]
print(a[1])
```

Assign your answer to the variable `q6_0` in the next cell. Do NOT put quotes around your answer. Your answer should be a numeric value.

```python
# ANSWER CELL
q6_0 = ...
q6_0
```

```python
grader.check("q6.0")
```

There are essential built-in functions that apply to data structures in Python. The operations in the table below are widely supported by most data structure types and form the foundation of data manipulation in Python.

In the table, `s` is a data structure, such as a list, and `i` is an integer.

| Operation      | Result                                                                      | 
| :------------- | :-------------------------------------------------------------------------- |
| `s[i]`         | Index and retrieve the item at position `i` of `s`                          |
| `len(s)`       | Return the length (number of items) in `s`                                  |
| `min(s)`       | Return the smallest or lowest item in `s`                                   |
| `max(s)`       | Return the largest or highest item in `s`                                   |

<font color='red'>**Question 6.1.**</font> Write a Python expression that extracts the maximum element in `a` and the minimum element in `a` and then divides their sum by the length of `a`. Assign this expression to `q6_1`. (0.5 pts)

```python
# ANSWER CELL
a = [383, 165, 273, 112, 937, 11, 64, 112, 455]
q6_1 = ...
q6_1
```

```python
grader.check("q6.1")
```

Another essential data structure for scientific and numerical applications is the NumPy array, which is part of the NumPy module. NumPy is short for "Numerical Python". 

In order to use NumPy, we need to import it first. A conventional way to import it is to use `np` as a shortened name: `import numpy as np`.

To define an array in Python, you could use the `np.array()` function. The items are put inside square brackets `np.array([...])` or a second set of parentheses `np.array((...))`, and are separated by commas `,`.

**Example:**

```python
>>> w = np.array([1, 4, 3, 12])
```

Arrays can have multiple rows and columns. In NumPy, attributes are properties of NumPy arrays that provide information about the array's shape, size, dimension, and so on. There are many attributes, which you can read about in the documentation [here](https://numpy.org/doc/stable/reference/generated/numpy.ndarray.html). Some of the common attributes are:
* [`array_name.shape`](https://numpy.org/doc/stable/reference/generated/numpy.ndarray.shape.html): returns the shape of `array_name`
> returns `(n_rows, n_columns)`, where the first element is the number of rows in `array_name` and the second element is the number of columns in `array_name` 
* [`array_name.size`](https://numpy.org/doc/stable/reference/generated/numpy.ndarray.size.html#numpy.ndarray.size): returns the total number of elements in `array_name` 
> number of rows times number of columns
* [`array_name.ndim`](https://numpy.org/doc/stable/reference/generated/numpy.ndarray.ndim.html#numpy.ndarray.ndim): returns the number of dimensions of `array_name`

In addition to `np.array()`, there are a number of other functions for creating arrays. Very often we would like to define arrays that have a pattern. For instance, we may wish to create the array `days = [1, 2, 3, ..., 365]`. It would be very cumbersome to type all the values of `days` into Python. For generating arrays that have a pattern, there are several helpful functions in NumPy.

For generating arrays that are in **order and evenly spaced**, it is useful to use the `np.arange()` function in NumPy:

```python
np.arange(start, stop, step)
```

where:

* `start`: a number specifying at which value to start the array. This is an optional argument, and thus, if not specified, it will take the default value 0.
* `stop`: a number specifying at which value to stop the array (excluded). This is a required argument.
* `step`: a number specifying the increment in the values of the array (`step` $\neq 0$). This is an optional argument, and thus, if not specified, it will take the default value 1. The `step` can be positive or negative and the value in index `i` will simply be `start + step*i`

**Example:**

```python
>>> days = np.arange(1, 366)
```

For generating arrays that have **a certain number of evenly spaced points between a start and end value**, it is useful to use the `np.linspace()` function in NumPy:

```python
np.linspace(start, stop, num)
```

where:

* `start`: a number specifying at which value to start the array. Here, this is a required argument.
* `stop`: a number specifying at which value to stop the array. This is by default **included** in the array, unlike `np.arange()`. This is a required argument.
* `num`: an integer number specifying how many values to generate between `start` and `stop`, including `start` and `stop`. This is an optional argument, and thus, if not specified, it will take the default value 50.

**Example:**

```python
>>> days = np.linspace(1, 365, 365)
```

Similar to strings and lists, you can you can index a NumPy array.

<font color='red'>**Question 6.2.**</font> What will be displayed after running this Python code? Feel free to run it to check your answer, but try to think it through first. (1 pt)

```python
x1 = np.arange(10, 20, 1)
x2 = np.linspace(0, 10, 3)
print(x1[-1] + x2[1])
```

Assign your answer to the variable `q6_2` in the next cell. Do NOT put quotes around your answer. Your answer should be a numeric value.

```python
# ANSWER CELL
q6_2 = ...
q6_2
```

```python
grader.check("q6.2")
```

### You're done with this Lab!

**Important submission information:** After completing the assignment, click on the Save icon from the Tool Bar &nbsp;<i class="fa fa-save" style="font-size:16px;"></i>&nbsp;. After saving your notebook, **run the cell with** `grader.check_all()` and confirm that you pass the same tests as in the notebook. Then, **run the final cell** `grader.export()` and click the link to download the zip file. Finally, go to Gradescope and submit the zip file to the corresponding assignment. 

**Once you have submitted, stay on the Gradescope page to confirm that you pass the same tests as in the notebook.**

```python
%matplotlib inline
img = mpimg.imread('resources/animal.jpg')
imgplot = plt.imshow(img)
imgplot.axes.get_xaxis().set_visible(False)
imgplot.axes.get_yaxis().set_visible(False)
print("Congratulations on finishing your first lab!")
plt.show()
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Make sure you submit the .zip file to Gradescope.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q1.0 (points: 0.5)
```
>>> assert len(q1_0) == 1
>>> assert isinstance(q1_0, str)
```
failure_message: Incorrect answer format. Make sure you only have the letter of your answer choice in quotes.
```
>>> assert get_hash(q1_0.upper()) == '0d61f8370cad1d412f80b84d143e1257'
```
failure_message: Incorrect answer.


Tests for q2.0 (points: 0.5)
```
>>> assert len(q2_0) == 1
>>> assert isinstance(q2_0, str)
```
failure_message: Incorrect answer format. Make sure you only have the letter of your answer choice in quotes.
```
>>> assert get_hash(q2_0.upper()) == '0d61f8370cad1d412f80b84d143e1257'
```
failure_message: Incorrect answer.


Tests for q2.1 (points: 0.5)
```
>>> assert get_hash(round(q2_1, 10)) == '012bd0771172d790b0c88107c9743442'
```
failure_message: Incorrect answer.


Tests for q2.2 (points: 0.5)
```
>>> assert get_hash(round(x, 1)) == 'c81e728d9d4c2f636f067f89cc14862c'
>>> assert get_hash(round(y, 1)) == 'e4da3b7fbbce2345d7772b0674a318d5'
```
failure_message: Incorrect x and y. Make sure x and y are defined based on the given values.
```
>>> assert get_hash(round(q2_2, 2)) == '027e97349ab774d836e6bde002c96d80'
```
failure_message: Incorrect answer for q2_2.


Tests for q3.0 (points: 0.5)
```
>>> assert len(q3_0) == 1
>>> assert isinstance(q3_0, str)
```
failure_message: Incorrect answer format. Make sure you only have the letter of your answer choice in quotes.
```
>>> assert get_hash(q3_0.upper()) == '7fc56270e7a70fa81a5935b72eacbe29'
```
failure_message: Incorrect answer.


Tests for q3.1 (points: 0.5)
```
>>> assert len(q3_1) == 1
>>> assert isinstance(q3_1, str)
```
failure_message: Incorrect answer format. Make sure you only have the letter of your answer choice in quotes.
```
>>> assert get_hash(q3_1.upper()) == '3a3ea00cfc35332cedf6e5e9a32e94da'
```
failure_message: Incorrect answer.


Tests for q3.2 (points: 0.5)
```
>>> assert len(q3_2) == 1
>>> assert isinstance(q3_2, str)
```
failure_message: Incorrect answer format. Make sure you only have the letter of your answer choice in quotes.
```
>>> assert get_hash(q3_2.upper()) == '3a3ea00cfc35332cedf6e5e9a32e94da'
```
failure_message: Incorrect answer.


Tests for q3.3 (points: 1)
```
>>> assert isinstance(q3_3, tuple)
>>> assert all((isinstance(option, str) for option in q3_3))
```
failure_message: Incorrect answer format. Make sure you only have the letter of your answer choice in quotes and answers are separated by commas.
```
>>> assert get_hash('A' in q3_3 or 'a' in q3_3) == 'f8320b26d30ab433c5a54546d21f414c'
```
failure_message: Incorrect answer.
```
>>> assert get_hash('B' in q3_3 or 'b' in q3_3) == 'f8320b26d30ab433c5a54546d21f414c'
```
failure_message: Incorrect answer.
```
>>> assert get_hash('C' in q3_3 or 'c' in q3_3) == 'f827cf462f62848df37c5e1e94a4da74'
```
failure_message: Incorrect answer. Can variable names begin with a digit?
```
>>> assert get_hash('D' in q3_3 or 'd' in q3_3) == 'f827cf462f62848df37c5e1e94a4da74'
```
failure_message: Incorrect answer. Can variable names begin with symbols?
```
>>> assert get_hash('E' in q3_3 or 'e' in q3_3) == 'f827cf462f62848df37c5e1e94a4da74'
```
failure_message: Incorrect answer. Can variable names begin with symbols?
```
>>> assert get_hash('F' in q3_3 or 'f' in q3_3) == 'f827cf462f62848df37c5e1e94a4da74'
```
failure_message: Incorrect answer. Can variable names include spaces?
```
>>> assert get_hash('G' in q3_3 or 'g' in q3_3) == 'f8320b26d30ab433c5a54546d21f414c'
```
failure_message: Incorrect answer.


Tests for q4.0 (points: 0.5)
```
>>> assert len(q4_0) == 1
>>> assert isinstance(q4_0, str)
```
failure_message: Incorrect answer format. Make sure you only have the letter of your answer choice in quotes.
```
>>> assert get_hash(q4_0.upper()) == '7fc56270e7a70fa81a5935b72eacbe29'
```
failure_message: Incorrect answer.


Tests for q4.1 (points: 0.5)
```
>>> assert len(q4_1) == 1
>>> assert isinstance(q4_1, str)
```
failure_message: Incorrect answer format. Make sure you only have the letter of your answer choice in quotes.
```
>>> assert get_hash(q4_1.upper()) == '9d5ed678fe57bcca610140957afab571'
```
failure_message: Incorrect answer.


Tests for q4.2 (points: 0.5)
```
>>> assert isinstance(q4_2, int)
```
failure_message: Incorrect answer format. Your answer should be an integer.
```
>>> assert get_hash(int(q4_2)) == 'c81e728d9d4c2f636f067f89cc14862c'
```
failure_message: Incorrect answer.


Tests for q4.3 (points: 0.5)
```
>>> assert len(q4_3) == 1
>>> assert isinstance(q4_3, str)
```
failure_message: Incorrect answer format. Make sure you only have the letter of your answer choice in quotes.
```
>>> assert get_hash(q4_3.upper()) == '7fc56270e7a70fa81a5935b72eacbe29'
```
failure_message: Incorrect answer.


Tests for q5.0 (points: 0.5)
```
>>> assert len(q5_0) == 1
>>> assert isinstance(q5_0, str)
```
failure_message: Incorrect answer format. Make sure you only have the letter of your answer choice in quotes.
```
>>> assert get_hash(q5_0.upper()) == 'f623e75af30e62bbd73d6df5b50bb7b5'
```
failure_message: Incorrect answer.


Tests for q5.1 (points: 1)
```
>>> assert isinstance(q5_1, tuple)
>>> assert all((isinstance(option, str) for option in q3_3))
```
failure_message: Incorrect answer format. Make sure you only have the letter of your answer choice in quotes and answers are separated by commas.
```
>>> assert get_hash('A' in q5_1 or 'a' in q5_1) == 'f827cf462f62848df37c5e1e94a4da74'
```
failure_message: Incorrect answer.
```
>>> assert get_hash('B' in q5_1 or 'b' in q5_1) == 'f8320b26d30ab433c5a54546d21f414c'
```
failure_message: Incorrect answer. You cannot concatenate strings and numeric data.
```
>>> assert get_hash('C' in q5_1 or 'c' in q5_1) == 'f8320b26d30ab433c5a54546d21f414c'
```
failure_message: Incorrect answer. Formatted strings should start with f.
```
>>> assert get_hash('D' in q5_1 or 'd' in q5_1) == 'f827cf462f62848df37c5e1e94a4da74'
```
failure_message: Incorrect answer.
```
>>> assert get_hash('E' in q5_1 or 'e' in q5_1) == 'f8320b26d30ab433c5a54546d21f414c'
```
failure_message: Incorrect answer. Variables should be inside {}.


Tests for q6.0 (points: 0.5)
```
>>> assert isinstance(q6_0, int)
```
failure_message: Incorrect answer format. Your answer should be an integer.
```
>>> assert get_hash(int(q6_0)) == '9766527f2b5d3e95d4a733fcfb77bd7e'
```
failure_message: Incorrect answer. Indexing in Python starts at 0.


Tests for q6.1 (points: 0.5)
```
>>> assert isinstance(q6_1, float)
```
failure_message: Incorrect answer format. Your answer should be a float.
```
>>> assert get_hash(round(q6_1, 10)) == '6d025b8f42889608be47f9567281fdd8'
```
failure_message: Incorrect answer. The entire sum should be divided by the length.


Tests for q6.2 (points: 1)
```
>>> assert isinstance(q6_2, int)
```
failure_message: Incorrect answer format. Your answer should be an integer.
```
>>> assert get_hash(int(q6_2)) == '1ff1de774005f8da13f42943881c655f'
```
failure_message: Incorrect answer. Indexing in Python starts at 0.
